<a href="https://colab.research.google.com/github/niusha73/CE880_Lab_Task/blob/main/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

---
# Welcome to CE880
### This is your week-1 : Problem 1 notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/).


---
## Question 1:

Write a function for Pythagorean theorem, or Pythagoras' theorem. Pythagorean is a fundamental relation in Euclidean geometry among the three sides of a right triangle. It states that the area of the square whose side is the hypotenuse (the side opposite the right angle) is equal to the sum of the areas of the squares on the other two sides. This theorem can be written as an equation relating the lengths of the legs a, b and the hypotenuse c, often called the Pythagorean equation:

$a^2 + b^2 = c^2$

Please write a function (code)to calculate the length of an unknown side of a right triangle, given two sides are known.

<!-- grade_id: myPythagorean_ans, points: 0, grade: false, solution: true -->
```python
import math
def myPythagorean(a,b):
    """write a function (code)to calculate the length of an unknown side
    of a right triangle, given two sides are known.
    """

    if a == 0 or b == 0:
      print("please entar valid number" )
    else:
      c = math.sqrt (a**2 + b ** 2)
      print ("right triangle is ", c)
    #raise NotImplementedError()
```

```python
myPythagorean(10,14)
```

Output:
```
right triangle is  17.204650534085253
```

<!-- grade_id: myPythagorean_test, points: 1, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
assert myPythagorean(3,2) == 3.605551275463989
assert myPythagorean(7,3) == 7.615773105863909
assert myPythagorean(0,2) == 0
```

Output:
```
right triangle is  3.605551275463989
AssertionError                            Traceback (most recent call last)
<ipython-input-13-df915bc8d9cc> in <cell line: 2>()
      1 """Check that squares returns the correct output for several inputs"""
----> 2 assert myPythagorean(3,2) == 3.605551275463989
      3 assert myPythagorean(7,3) == 7.615773105863909
      4 assert myPythagorean(0,2) == 0

AssertionError:
```

---

## Question 2:

A `trapezoid` is a quadrilateral with two sides parallel. The `trapezoid` is equivalent to the British definition of the `trapezium`. An isosceles `trapezoid` is a `trapezoid` in which the base angles are equal so

<img src="https://www.w3resource.com/w3r_images/python-math-image-exercise-3.svg"
     alt="Markdown Monster icon"
     style="float: left; margin-right: 10px;" />
     
Please write a function (code)to calculate the area of trapezoid.

<!-- grade_id: areaofTrap_ans, points: 0, grade: false, solution: true -->
```python
def areaofTrap(base1, base2, height):
    """Compute the sum of the squares of numbers from 1 to n."""
    A = 0.5 * (base1 + base2) * height
    print("area of trap is ", A)
    #raise NotImplementedError()
```

```python
areaofTrap(5, 3, 6)
```

Output:
```
area of trap is  24.0
```

<!-- grade_id: areaofTrap_test, points: 1, grade: true, solution: false -->
```python
"""Check that areaofTrap returns the correct answer for various inputs."""
assert areaofTrap(10, 10, 20) == 200.0
assert areaofTrap(5, 10, 20) == 150.0
assert areaofTrap(5, 3, 5) == 20.0
```

---
## Question 3:

 Write a Python program to reverse a string. Example: Sample String: `1234xyz` Expected Output: `zyx4321`

<!-- grade_id: reversestring_ans, points: 0, grade: false, solution: true -->
```python
def reverse_string(str1):
    str1 = input('Please enter your string: ')
    return str1 [::-1]

    #raise NotImplementedError()
```

```python
reverse_string('ChatGPT')

```

Output:
```
Please enter your string: ChatGPT
'TPGtahC'
```

<!-- grade_id: reversestring_test, points: 1, grade: true, solution: false -->
```python
"""Check that areaofTrap returns the correct answer for various inputs."""
assert reverse_string('Humanitarian') == 'nairatinamuH'
assert reverse_string('ChatGPT') == 'TPGtahC'
assert reverse_string('Data Science') == 'ecneicS ataD'
assert reverse_string('CE880') == '088EC'
```

## Question 4:
Write a Python program to print the even numbers from a given list and in ascending order.

Sample List : `[1, 2, 3, 4, 5, 6, 7, 8, 9]` Expected Result : `[2, 4, 6, 8]`

<!-- grade_id: even_num_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np
def even_num(lst_1):
    """Write a Python program to print the even numbers from a given list and in ascending order."""
    lst_1.sort()
    index = 0
    while index< len(lst_1):
      number = lst_1 [index]
      if number % 2 == 0:
        print(number)
      index +=1


#raise NotImplementedError()
```

```python
even_num([2, 3, 14, 17, 8, 19, 11, 14])
```

Output:
```
2
8
14
14
```

<!-- grade_id: odd_num_test, points: 1, grade: true, solution: false -->
```python
assert even_num([1, 2, 3, 4, 6, 7, 8, 9]) == [2, 4, 6, 8]
assert even_num([2, 3, 4, 7, 8, 9, 11, 14]) == [2, 4, 8, 14]
assert even_num([2, 3, 14, 17, 8, 19, 11, 14]) == [2, 8, 14]
```

Output:
```
2
4
6
8
AssertionError                            Traceback (most recent call last)
<ipython-input-20-c69ddecc5420> in <cell line: 1>()
----> 1 assert even_num([1, 2, 3, 4, 6, 7, 8, 9]) == [2, 4, 6, 8]
      2 assert even_num([2, 3, 4, 7, 8, 9, 11, 14]) == [2, 4, 8, 14]
      3 assert even_num([2, 3, 14, 17, 8, 19, 11, 14]) == [2, 8, 14]

AssertionError:
```

## Question 5:
Write a Python function to create and print a list where the values are square of numbers between `1` and `n`, where `n` is user defined number: Example `sqr_values(4) == [1, 4, 9, 16]`

<!-- grade_id: sqr_values_ans, points: 0, grade: false, solution: true -->
```python
def sqr_values(n):

  my_list = []
  i = 1
  for i in range (1, n+1):

    my_list.append(i ** 2)
    i+= 1
  return my_list



    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
sqr_values(8)
```

Output:
```
[1, 4, 9, 16, 25, 36, 49, 64]
```

<!-- grade_id: sqr_values_test, points: 1, grade: true, solution: false -->
```python
"""Check that sqr_values returns the correct answer for various inputs."""
assert sqr_values(4) == [1, 4, 9, 16]
assert sqr_values(8) == [1, 4, 9, 16, 25, 36, 49, 64]
assert sqr_values(9) == [1, 4, 9, 16, 25, 36, 49, 64, 81]
assert sqr_values(11) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121]
```

Output:
```
0
1
4
9
<ipython-input-35-6ec84a6f7fdd> in <cell line: 2>()
      1 """Check that sqr_values returns the correct answer for various inputs."""
----> 2 assert sqr_values(4) == [1, 4, 9, 16]
      3 assert sqr_values(8) == [1, 4, 9, 16, 25, 36, 49, 64]
      4 assert sqr_values(9) == [1, 4, 9, 16, 25, 36, 49, 64, 81]
      5 assert sqr_values(11) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121]

<ipython-input-33-bc4769d7e931> in sqr_values(n)
     13 
     14     # YOUR CODE HERE
---> 15     raise NotImplementedError()

NotImplementedError:
```

## Question 6:

Write a Python program to find all the duplicate characters which are similar to each other and join them.


<!-- grade_id: comb_ans, points: 0, grade: false, solution: true -->
```python
def findDupChar(input):
    """Write a Python program to find all the duplicate characters which are similar to each other."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
findDupChar('Essex University')
```

<!-- grade_id: comb_test, points: 1, grade: true, solution: false -->
```python
assert findDupChar('Happyy') == 'p y'
assert findDupChar('Essex University') == 's e i'
```

```python

```